Cortical Plasticity is associated with Blood-Brain-Barrier Modulation

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public Review):

The goal of the current study was to evaluate the effect of neuronal activity on blood-brain barrier permeability in the healthy brain, and to determine whether changes in BBB dynamics play a role in cortical plasticity. The authors used a variety of well-validated approaches to first demonstrate that limb stimulation increases BBB permeability. Using in vivo-electrophysiology and pharmacological approaches, the authors demonstrate that albumin is sufficient to induce cortical potentiation and that BBB transporters are necessary for stimulus-induced potentiation. The authors include a transcriptional analysis and differential expression of genes associated with plasticity, TGF-beta signaling, and extracellular matrix were observed following stimulation. Overall, the results obtained in rodents are compelling and support the authors' conclusions that neuronal activity modulates the BBB in the healthy brain and that mechanisms downstream of BBB permeability changes play a role in stimulus-evoked plasticity. These findings were further supported with fMRI and BBB permeability measurements performed in healthy human subjects performing a simple sensorimotor task. There is literature to suggest that there are sex differences in BBB dysfunction in pathophysiological conditions and the authors have acknowledged the use of only males as a minor limitation of the study that should be addressed in the future. Future studies should also test whether the upregulation of OAT3 plays a role in cortical plasticity observed following stimulation. Overall, this study provides novel insights into how neurovascular coupling, BBB permeability, and plasticity interact in the healthy brain.

Reviewer #2 (Public Review):

Summary:

This study builds upon previous work that demonstrated that brain injury results in leakage of albumin across the blood brain barrier, resulting in activation of TGF-beta in astrocytes. Consequently, this leads to decreased glutamate uptake, reduced buffering of extracellular potassium and hyperexcitability. This study asks whether such a process can play a physiological role in cortical plasticity. They first show that stimulation of a forelimb for 30 minutes in a rat results in leakage of the blood brain barrier and extravasation of albumin on the contralateral but not ipsilateral cortex. The authors propose that the leakage is dependent upon neuronal excitability and is associated with an enhancement of excitatory transmission. Inhibiting the transport of albumin or the activation of TGF-beta prevents the enhancement of excitatory transmission. In addition, gene expression associated with TGF-beta activation, synaptic plasticity and extracellular matrix are enhanced on the "stimulated" hemisphere. That this may translate to humans is demonstrated by a break down in the blood brain barrier following activation of brain areas through a motor task.

Strengths:

This study is novel and the results are potentially important as they demonstrate an unexpected break down of the blood brain barrier with physiological activity and this may serve a physiological purpose, affecting synaptic plasticity.

The strengths of the study are:

(1) The use of an in vivo model with multiple methods to investigate the blood brain barrier response to a forelimb stimulation.

(2) The determination of a potential functional role for the observed leakage of the blood brain barrier from both a genetic and electrophysiological view point

(3) The demonstration that inhibiting different points in the putative pathway from activation of the cortex to transport of albumin and activation of the TGF-beta pathway, the effect on synaptic enhancement could be prevented. (4) Preliminary experiments demonstrating a similar observation of activity dependent break down of the blood brain barrier in humans.

Weaknesses:

The authors adequately addressed most of my points. A few remain:

(1) Although the reviewers have addressed the possible effects of anaesthesia on neuro-vascular coupling. They have not mentioned or addressed the possible effects of ketamine (an NMDA receptor antagonist) on synaptic plasticity. Indeed, the low percentage of SEP increase following potentiation (10-20%) could perhaps be explained by partial block of NMDA receptors by ketamine.

We agree and apologize for this oversight. This important issue is now addressed in the Discussion.

“Notably, the antagonistic effect of ketamine on NMDA receptors might attenuate the magnitude of SEP potentiation recorded in our experiments (Anis et al., 1983; Salt et al., 1988).”

(2) The experimental paradigms remain unclear to me. Now, it appears that drugs are applied for 50 minutes and that the stimulation occurs during the "washout period". The more conventional approach would be to have the drug application during the stimulation period to determine if the drugs occlude or enhance the effects of stimulation and then washout the drugs. The problem is that drugs variably washout at different rates depending upon their lipid solubility.

We agree that the more conventional approach would have been to continue applying the drug throughout the experiment and that differential rates of washout may add variability to our experiments. However, despite this limitation, within each treatment group we found that the SEP response at 50 minutes (immediately after the drug application window) does not differ from SEP response at 80 minutes (after 30 minutes of stimulation and washout) [Figure 3H&G]. This suggests that the drug effects were still present despite terminating drug application and performing potentiation-inducing stimulation. Moreover, our analysis showed that animals within each treatment group (except AP5) had similar SEP responses with little intra-group variability.

(3) It is still not clear to what extent the experimenters and those doing the analysis were blinded to group. If one or both were blind to group, then please put this in the methods.

Thank you for this comment. We revised the Methods section to clearly confirm that data was collected and analyzed blindly.

Reviewer #3 (Public Review):

Summary:

This study used prolonged stimulation of a limb to examine possible plasticity in somatosensory evoked potentials induced by the stimulation. They also studied the extent that the blood brain barrier (BBB) was opened by the prolonged stimulation and whether that played a role in the plasticity. They found that there was potentiation of the amplitude and area under the curve of the evoked potential after prolonged stimulation and this was long-lasting (>5 hrs). They also implicated extravasation of serum albumin, caveolae-mediated transcytosis, and TGFb signalling, as well as neuronal activity and upregulation of PSD95. Transcriptomics was done and implicated plasticity related genes in the changes after prolonged stimulation, but not proteins associated with the BBB or inflammation. Next, they address the application to humans using a squeeze ball task. They imaged the brain and suggest that the hand activity led to an increased permeability of the vessels, suggesting modulation of the BBB.

Strengths:

The strengths of the paper are the novelty of the idea that stimulation of the limb can induce cortical plasticity in a normal condition, and it involves opening of the BBB with albumin entry. In addition, there are many datasets and both rat and human data.

Weaknesses:

The conclusions are not compelling however because of a lack of explanation of methods.

In the revised paper, we added a section titled ‘study design’ that presents an overview of the experimental approach.

The explanation of why prolonged stimulation in the rat was considered relevant to normal conditions should be as clear in the paper as it is in the rebuttal.

We added a new paragraph to the Discussion section explaining this point as we did in the rebuttal:

“Our animal experiments show that a 30 min limb stimulation (at 6Hz and 2mA) increases cross-BBB influx, while a 1 min stimulation (of similar frequency and magnitude) does not. We believe that both types of stimulations fall within the physiological range because our continuous electrophysiological recordings showed no signs of epileptiform or otherwise pathological activity. Moreover, the recorded SEP levels were similar to those reported in previous physiological LTP studies in rats (Eckert & Abraham, 2010; Han et al., 2015; Mégevand et al., 2009) and humans (McGregor et al., 2016). In humans, skill acquisition often involves motor training sessions that last ≥30 minutes (Bengtsson et al., 2005; Classen et al., 1998) and result in physiological plasticity of sensory and motor systems (Classen et al., 1998; Draganski et al., 2004; Sagi et al., 2012). Hence, the experimental task in our human study (30 minutes of repetitive squeezing of an elastic stress-ball) is likely to represent physiological activity, with neuronal activation in primarily motor and sensory areas (Halder et al., 2005). Future human and animal studies are needed to explore the BBB modulating effects of additional stimulation protocols – with varying durations, frequencies, and magnitudes. Such studies may also elucidate the temporal and ultrastructural characteristics that differentiate between physiological and pathological BBB modulation. “

The authors need to ensure other aspects of the rebuttal are as clear in the paper as in the rebuttal too.

Thank you for this comment. This was addressed in the revised paper.

The only remaining concern that is significant is that it is hard to understand the figures.

Thank you for this comment. We revised the figures according to the reviewer’s recommendations. We hope that these changes increase the legibility of the figures.

Reviewer #3 (Recommendations For The Authors):

The manuscript is improved but there are still suggestions that do not appear to have been addressed. More experiments are not involved in addressing these concerns but one wants the paper to be clarified in terms of what was done.

Figures. Please use arrows to point to the effect that the reader should see. Please note what the main point is.

Major concerns:

Please add explanations, exact p values, and other revisions in the rebuttal to the paper.

Rebuttal explanations were added to the paper and p values appear in figure legends.

Fig 1d shows a seizure-like event which the authors don't think is a seizure because it lacks a depolarization ship. This explanation is not convincing because a LFP would not necessarily show a depolarization ship. Another argument of a discussion of the event as a seizure is warranted. Note that expanding the trace might also show it is unlike a seizure. Regarding the idea that 6Hz 2 mA stimuli for 30 min are physiological, the authors make three arguments which are not clear. First, no epileptiform activity was found, but in Fig. 1 it looks like a seizure occurred. Second, memory and skill acquisition in humans open involve a similar training duration - but what about 6Hz 2 mA?

Rats are known to rhythmically move their whiskers at frequencies ranging between 5 and 15 Hz (Mégevand et al., 2009). We agree that there is no clear way to justify the similarity between the experimental design in humans and rats. However, we believe that both paradigms (paw stimulation in rats and ball squeeze in humans) represent non-pathological input that we found to modulate barrier permeability. This argument was added to the discussion of the paper:

“We believe that both types of stimulations fall within the physiological range because in rats, activity between 515 Hz represents physiological rhythmic whisker movement during environment exploration (Mégevand et al., 2009).”

Seizures are typically induced in rats via direct tetanic stimulation of the brain (at 50 Hz and 0.3-2.5mA) or maximal electroshock test to the cornea (at 50 Hz and 150 mA) (Swinyard et al., 1952). We, therefore, assert that the activity we observe represents physiological responses and not seizures. This argument is beyond the scope of the current paper.

Please note a limitation is that the high level of serum albumin is unlikely to be physiological but may not have been as high in the animal because of the low diffusion rate and degradation (please add the refs in the rebuttal).

Thank you, we added the following to the Results section:

“The relatively high concentration of albumin was chosen to account for factors that lower its effective tissue concentration such as its low diffusion rate and its likelihood to encounter a degradation site or a cross-BBB efflux transporter (Tao & Nicholson, 1996; Zhang & Pardridge, 2001).”

Fig. 1.

Please consider a box in b to show where the expanded traces in the lower row came from.

Thank you for the suggestion. We added lines indicating where the trace excerpts were taken from.

c. Please use arrows to point to the parts that the authors want the reader to note. In the legend, explain what t is, and delta HbT.

Thank you. We implemented this suggestion.

d. It is not clear what the double-sided arrows are meant to show compared to the arrow without two sides.

We replaced the two-headed arrow with two single ones.

e. Please explain what the upward lines at the top signify. What does the red asterisk mean?

Thank you. We implemented this suggestion.

f. Is the reader supposed to note the yellow area? Please make it with an arrow or circle if so.

Thank you, we added a white circle to mark the area of tracer accumulation.

g. Please explain what the permeability index is or reference the part of the paper that does.

Further to this suggestion, we added a refence to the appropriate methods section to the legend.

h. Please use arrows to point to the area of interest.

Thank you. We implemented this suggestion.

m-n. Please mark areas of interest with arrows. m. the top right two images are unclear. I suggest making them say ipsi inset and contra inset instead of using asterisks.

Thank you. We added the ipsi and contra labels to panels in m. The images in panel n represent a phenomenon with no particular region of interest, but rather peri-vascular tracer accumulation along the entire depicted blood vessel. We clarified that panel n represents a separate experiment than panel m: “n. In an animal injected with both EB and NaFlu post stimulation, fluorescence imaging shows extravascular accumulation of both tracers along a cortical small vessel in the stimulated hemisphere.”

Figure 2.

(2) a. Middle. What are the vertical lines at the top? The rebuttal states that was explained in the revised legends but I don't see it.

Our apologies. We now included an explanation that “an excerpt of the stimulation trace is shown above the middle LFP trace”.

c and d are very different field potentials in shape and therefore hard to compare. The rebuttal addresses this but the explanation is not in the revised text.

We agree that there is variability in SEP responses between animals. We now added a statement acknowledging this in the methods section: “To overcome potential variability in SEP morphology between animals (Mégevand et al., 2009), each animal’s plasticity measures (max amplitude and AUC of post stimulation SEP) were compared to the same measures at baseline.”

In d, it is not clear there is potentiation because the traces are not aligned.

All panels depicting SEP traces represent raw data with no alignment. The shift observed in panel d exemplifies why we compare post-stimulation parameters of max amplitude and area under curve to baseline in each animal.

Exact P values are said to have been added in the rebuttal but they were not.

Exact P values appear in Figure legends.

(3) b. Use arrows to mark the area of interest.

Thank you. We added a white circle to mark the area of tracer accumulation similar to Figure 1f.

d. Why is there an oscillation superimposed on all traces except CNQX?

We agree that this is an interesting question. Future studies should determine the source of this SEP pattern.

(4) What does the line and the number 2 mean? How were data normalized? What was counted? What area of cortex?

The number 2 refers to the scale bar line, meaning a log fold change of 2 reflects the size of the scale bar line.

The plot shows the log fold change against the mean count of each gene in the contralateral somatosensory cortex between 1 and 24 hours after stimulation.

The x axis title was changed to “mean expression” and the legend was modified to:

“Scatter plot of gene expression from RNA-seq in the contralateral somatosensory cortex 24 vs. 1 h after 30 min stimulation. The y axis represents the log fold change, and the x axis represents the mean expression levels (see methods, RNA Sequencing & Bioinformatics). Blue dots indicate statistically significant differentially expressed genes (DEGs) by Wald Test (n=8 rats per group).”

How were the pericytes, smooth muscle cells, ,etc. distinguished?

This was explained under Methods->RNA Sequencing & Bioinformatics: “Analysis of cell-specific and vascular zonation genes was performed as described (Vanlandewijck et al., 2018), using the database provided in (http://betsholtzlab.org/VascularSingleCells/database.html).”

What were the chi square statistics? If there were cells used instead of rats, please justify.

Thank you. The legend was expanded to include the following:

“The contralateral somatosensory cortex was found to have a significantly higher number of DEGs related to synaptic plasticity, than the ipsilateral side (***p<0.001, Chi-square).”

(5) b. what do the icons mean?

We agree that the icons were confusing. We simplified this panel to just show when participants were asked to squeeze the ball (black icon). This explanation was added to the Figure legend.

Abbreviations?

Abbreviations of MRI protocols were added to the figure legend for clarity.

In c-e what are the units of measure? Fold-change?

The units represent t-statistics values for each voxel. The label ‘t-statistic’ was added to the figure.

What are the white Iines, + and - signs?

The white lines point to voxels of highest activation (t-statistic). This was added to the legend.

And these are not +/- signs these are voxels with significant activation which only appear similar.

f. Please explain f and g for clarity.

Thank you. The explanation was modified for added clarity.

Supplemental Fig. 4.

Original question: If ipsilateral and contralateral showed many changes why do the authors think the effects were only contralateral?

The authors replied: Our gene analysis was designed to complement our in vivo and histological findings, by assessing the magnitude of change in differentially expressed genes (DEGs). This analysis showed that: (1) the hemisphere contralateral to the stimulus has significantly more DEGs than the ipsilateral hemisphere; and (2) the DEGs were related to synaptic plasticity and TGF-b signaling. These findings strengthen the hypothesis raised by our in vivo and histological experiments.

Could the authors clarify the answer to the question in the text?

Thank you. This section was added to the Discussion.

Papers referenced in this letter:

Anis, N. A., Berry, S. C., Burton, N. R., & Lodge, D. (1983). The dissociative anaesthetics, ketamine and phencyclidine, selectively reduce excitation of central mammalian neurones by N-methyl-aspartate. British Journal of Pharmacology, 79(2), 565–575. hQps://doi.org/10.1111/j.1476-5381.1983.tb11031.x

Bengtsson, S. L., Nagy, Z., Skare, S., Forsman, L., Forssberg, H., & Ullén, F. (2005). Extensive piano practicing has regionally specific effects on white matter development. Nature Neuroscience, 8(9), 1148–1150. hQps://doi.org/10.1038/nn1516

Classen, J., Liepert, J., Wise, S. P., Hallett, M., & Cohen, L. G. (1998). Rapid plasticity of human cortical movement representation induced by practice. Journal of Neurophysiology, 79(2), 1117–1123. hQps://doi.org/10.1152/JN.1998.79.2.1117/ASSET/IMAGES/LARGE/JNP.JA47F4.JPEG

Draganski, B., Gaser, C., Busch, V., Schuierer, G., Bogdahn, U., & May, A. (2004). Changes in grey matter induced by training. Nature, 427(6972), 311–312. hQps://doi.org/10.1038/427311a

Eckert, M. J., & Abraham, W. C. (2010). Physiological effects of enriched environment exposure and LTP induction in the hippocampus in vivo do not transfer faithfully to in vitro slices. Learning and Memory, 17(10), 480–484. hQps://doi.org/10.1101/lm.1822610

Halder, P., Sterr, A., Brem, S., Bucher, K., Kollias, S., & Brandeis, D. (2005). Electrophysiological evidence for cortical plasticity with movement repetition. European Journal of Neuroscience, 21(8), 2271–2277. hQps://doi.org/10.1111/J.1460-9568.2005.04045.X

Han, Y., Huang, M. De, Sun, M. L., Duan, S., & Yu, Y. Q. (2015). Long-term synaptic plasticity in rat barrel cortex. Cerebral Cortex, 25(9), 2741–2751. hQps://doi.org/10.1093/cercor/bhu071

McGregor, H. R., Cashaback, J. G. A., & Gribble, P. L. (2016). Functional Plasticity in Somatosensory Cortex Supports Motor Learning by Observing. Current Biology, 26(7), 921–927. hQps://doi.org/10.1016/j.cub.2016.01.064

Mégevand, P., Troncoso, E., Quairiaux, C., Muller, D., Michel, C. M., & Kiss, J. Z. (2009). Long-term plasticity in mouse sensorimotor circuits after rhythmic whisker stimulation. Journal of Neuroscience, 29(16), 5326– 5335. hQps://doi.org/10.1523/JNEUROSCI.5965-08.2009

Sagi, Y., Tavor, I., HofsteQer, S., Tzur-Moryosef, S., Blumenfeld-Katzir, T., & Assaf, Y. (2012). Learning in the Fast Lane: New Insights into Neuroplasticity. Neuron, 73(6), 1195–1203. hQps://doi.org/10.1016/j.neuron.2012.01.025

Salt, T. E., Wilson, D. G., & Prasad, S. K. (1988). Antagonism of N-methylaspartate and synapBc responses of neurones in the rat ventrobasal thalamus by ketamine and MK-801. British Journal of Pharmacology,

94(2), 443–448. hQps://doi.org/10.1111/j.1476-5381.1988.tb11546.x

Swinyard, E. A., Brown, W. C., & Goodman, L. S. (1952). Comparative assays of antiepileptic drugs in mice and rats. The Journal of Pharmacology and Experimental Therapeutics, 106(3), 319–330. hQp://jpet.aspetjournals.org/content/106/3/319.abstract

Tao, L., & Nicholson, C. (1996). Diffusion of albumins in rat cortical slices and relevance to volume transmission. Neuroscience, 75(3), 839–847. hQps://doi.org/10.1016/0306-4522(96)00303-X

Vanlandewijck, M., He, L., Mäe, M. A., Andrae, J., Ando, K., Del Gaudio, F., Nahar, K., Lebouvier, T., Laviña, B.,

Gouveia, L., Sun, Y., Raschperger, E., Räsänen, M., Zarb, Y., Mochizuki, N., Keller, A., Lendahl, U., &

Betsholtz, C. (2018). A molecular atlas of cell types and zonation in the brain vasculature. Nature, 554(7693), 475–480. hQps://doi.org/10.1038/nature25739

Zhang, Y., & Pardridge, W. M. (2001). Mediated efflux of IgG molecules from brain to blood across the blood– brain barrier. Journal of Neuroimmunology, 114(1–2), 168–172. hQps://doi.org/10.1016/S01655728(01)00242-9

Changes from the earlier version to the revised version

Abstract

Summary

Brain microvessels possess the unique properties of a blood-brain barrier (BBB), tightly regulating the passage of molecules from the blood to the brain neuropil and vice versa. In models of brain injury, BBB dysfunction and the associated leakage of serum albumin to the neuropil have been shown to induce pathological plasticity, neuronal hyper-excitability, and seizures. The effect of neuronal activity on BBB function and whether it plays a role in plasticity in the healthy brain remain unclear. Here we show that neuronal activity induces modulation of microvascular permeability in the healthy brain and that it has a role in local network reorganization. Combining simultaneous electrophysiological recording and vascular imaging with transcriptomic analysis in rats, and functional and BBB-mapping MRI in human subjects we show that prolonged stimulation of the limb induces a focal increase in BBB permeability in the corresponding somatosensory cortex that is associated with long-term synaptic plasticity. We further show that the increased microvascular permeability depends on neuronal activity and involves caveolae-mediated transcytosis and transforming growth factor beta signaling. Our results reveal a role of BBB modulation in cortical plasticity in the healthy brain, highlighting the importance of neurovascular interactions for sensory experience and learning.

Introduction

The intricate components comprising the blood-brain barrier (BBB) regulate the bi-directional transport of molecules between the brain and the circulatory system and maintain the extracellular environment needed for normal function of the neurovascular unit (NVU) (Chow and Gu, 2015). Dysfunction of the BBB is common in many neurological disorders and is associated with impaired cross-BBB influx and efflux (Obermeier et al., 2013). For example, seizures were shown to increase cross-BBB influx of serum albumin – in a process partially mediated by excessive release of neuronal glutamate and subsequent activation of N-methyl-D-aspartate (NMDA) receptors (Vazana et al., 2016).

Emerging evidence suggests that cross-BBB influx/efflux may also change in response to physiological neuronal activity in the healthy brain: (1) whisker stimulation (1 hour, 2 Hz) was shown to increase BBB influx of insulin-like growth factor-I in rats (Nishijima et al., 2010); (2) physiological neuronal activity was shown to regulate BBB efflux transport in mice (Pulido et al., 2020); and (3) circadian rhythms were shown to alter BBB efflux in Drosophila (Zhang et al., 2018) and mice (Pulido et al., 2020). However, the mechanisms and the role of these physiological changes in BBB properties are not known.

Here, we sought to test whether physiological BBB modulation may involve mechanisms previously linked to pathological BBB changes, e.g., transcytosis of serum albumin (Ivens et al., 2007; Seiffert et al., 2004) and transforming growth factor β (TGF-β) signaling (David et al., 2009; Heinemann et al., 2012; Ivens et al., 2007; Weissberg et al., 2015). As TGF-β signaling has been linked to synaptic plasticity (Dahlmanns et al., 2022; Gradari et al., 2021; Patel and Weaver, 2021; Weissberg et al., 2015), we also explored the hypothesis that physiological BBB modulation may be associated with neuroplasticity. Lastly, we sought to provide the first evidence for physiological BBB modulation in humans.

Results

Stimulation increases BBB permeability

To test for changes in BBB permeability in response to stimulation we combined electrophysiological recordings with BBB imaging using intravital microscopy (Figure 1, see timeline of stimuli in Figure 1a). First, we localized vascular response to sensory stimulation using wide-field imaging of arterial diameter and monitoring change in total hemoglobin (HbT) signal during a 1 min stimulation of the limb (Figure 1b-d) (Bouchard et al., 2009). Fluorescent angiography before and immediately after 30 min of limb stimulation showed local extravasation of the tracer sodium fluorescein (NaFlu) around the responding blood vessels, indicative of BBB leakage (Figure 1f-g). Histological analysis confirmed the presence of NaFlu in the extravascular space around small vessels in the somatosensory cortex, contralateral (and not ipsilateral) to the stimulated limb (Figure 1m-n). Next, we intravenously injected a separate cohort of rats with either the albumin-binding dye Evans blue (EB) or albumin conjugated with Alexa Fluor 488 (Alexa488-Alb) without performing craniotomy. Brain extravasation of both dyes was found in the contralateral hemisphere, indicating the presence of serum albumin in the neuropil (Figure 1h-k). ELISA performed in cortical tissue at different time points after stimulation confirmed a higher concentration of albumin in the stimulated sensorimotor cortex early after stimulation (30 min) compared to non-stimulated sham-controls, that declined 4- and 24-hours post-stimulation (Figure 1l). As a positive control, brain tissues were also prepared from rats subjected to photothrombosis (PT)-induced ischemia (see methods). Albumin concentrations in the peri-ischemic cortex 24 hours following PT were significantly higher compared to sham and all time-points after stimulation (Figure S1d). These results suggest that the focal modulation of BBB permeability following stimulation is distinct from the extensive BBB dysfunction observed in injury models.

Limb stimulation modulates blood-brain barrier permeability.

Earlier version: a. The experimental paradigm in control animals (Ctrl, top panel) and animals that underwent a 30 min stimulation (Stim, bottom panel). Gray lines indicate 1 min test stimulations, lightning indicates the timing of BBB imaging (using injection of sodium fluorescein, NaFlu). b. Top: local field potential (LFP) trace from the corresponding region in L2/3 sensorimotor cortex before during (greyed area) and after 30 min limb stimulation. Bottom: 5 s excerpts of the above trace (left-to-right: before during and after stimulation). c. Image of the cortical window over the rat sensorimotor cortex, followed by total hemoglobin (HbT) concentration maps showing the evolution of the hemodynamic response during and after stimulation. Red rectangle marks the responding region magnified in d. d. Images of the responding arteriole (marked with red arrows) in the rat’s cortex before (left) during (middle) and after stimulation (right). e. The diameter of the responding arteriole during 1 min stimulation. Gray area corresponds to time of stimulation. f. NaFlu permeability maps before (left) and after 30 min stimulation (right) showing tracer accumulation around a responding arteriole. g. Permeability index is higher after stimulation compared to baseline (n=13 rats, mean ± SEM, Wilcoxon, p<0.001). h-i. Fluorescence images of cortical sections of stimulated rats injected with the albumin-binding dye Evans blue (h) and Alexa-488-Alb (i). j-k. Total fluorescence of EB (j) and Alexa488-Alb (k) after stimulation in the contralateral hemisphere compared to the ipsilateral hemisphere. (EB n=4 rats, 32 sections, nested t-test, p=0.0296; Alexa488-Alb n=5 rats, 20 sections, nested t-test, p=0.0229; mean ± SEM). l. Albumin concentration in the contralateral hemisphere relative to the ipsilateral in three different timepoints after stimulation, compared to sham stimulation. (0.5 and 4 h post-stimulation n=8, sham n=7, mean ± SEM, Kruskal-Wallis with FDR correction, p=0.0242, q=0.0406). m. Cortical sections of the area of limb representation from both hemispheres of a stimulated rat (left), and higher magnification images (right) n. Fluorescence image of a cortical small vessel in the stimulated hemisphere showing extravascular accumulation of EB and NaFlu.

Revised version: a. The experimental paradigm in control animals (Ctrl, top panel) and animals that underwent a 30 min stimulation (Stim, bottom panel). Gray lines indicate 1 min test stimulations, lightning indicates the timing of BBB imaging (using injection of sodium fluorescein, NaFlu). b. Top: local field potential (LFP) trace from the corresponding region in L2/3 sensorimotor cortex before, during (greyed area), and after 30 min limb stimulation. Bottom: 5 s excerpts of the above trace, at time point indicated by orange lines (left-to-right: before during and after stimulation). c. Image of the cortical window over the rat sensorimotor cortex, followed by the change in total hemoglobin (Λ[HbT]) concentration maps showing the evolution of the hemodynamic response during and after stimulation (t states the timepoint of the image; 120 s total: 30 s before and after a 60 s stimulation). Red rectangle marks the responding region magnified in d. d. Images of the responding arteriole (dilated parts marked with red arrows) in the rat’s cortex before (left) during (middle) and after stimulation (right). The red asterisk denotes the measured arteriole on panel e. e. The diameter of the responding arteriole during 1 min stimulation. Gray area corresponds to time of stimulation (the stimulation trace is shown above the grayed area). f. NaFlu permeability maps before (left) and after 30 min stimulation (right) showing tracer accumulation around a responding arteriole (marked with a white circle). g. Permeability index (see methods, under In Vivo imaging) is higher after stimulation compared to baseline (n=13 rats, mean ± SEM, Wilcoxon, p<0.001). h-i. Fluorescence images of cortical sections of stimulated rats injected with the albumin-binding dye Evans blue (h) and Alexa-488-Alb (i). White arrows point to areas with tracer accumulation. j-k. Total fluorescence of EB (j) and Alexa488-Alb (k) after stimulation in the contralateral hemisphere compared to the ipsilateral hemisphere. (EB n=4 rats, 32 sections, nested t-test, p=0.0296; Alexa488-Alb n=5 rats, 20 sections, nested t-test, p=0.0229; mean ± SEM). l. Albumin concentration in the contralateral hemisphere relative to the ipsilateral in three different timepoints after stimulation, compared to sham stimulation. (0.5 and 4 h post-stimulation n=8, sham n=7, mean ± SEM, Kruskal-Wallis with FDR correction, p=0.0242, q=0.0406). m. Cortical sections of the area of limb representation from both hemispheres of a stimulated rat (left), and higher magnification images (right) n. In an animal injected with both EB and NaFlu post stimulation, fluorescence imaging shows extravascular accumulation of both tracers along a cortical small vessel in the stimulated hemisphere.

BBB permeability is associated with synaptic potentiation

To study changes in brain activity during repetitive stimulation of the limb, we measured somatosensory evoked potentials (SEP) by averaging neuronal response to a 1 min test stimulation before and after the 30 min stimulation train (6 Hz, 2 mA). SEP amplitudes and area under the curve were increased in stimulated, but not in control cortices (Figure 2a-e). Potentiation lasted until the end of each experiment (Figure 2f) and was therefore indicative of long-term potentiation (LTP).

Stimulation and cortical perfusion of serum albumin induce LTP.

Earlier version: a. Top: LFP trace from the rat L2/3 sensorimotor cortex before (left) during (middle) and after (right) test stimulation (1 min, 6 Hz, 2 mA). Bottom: LFP trace from the same rat before (left) during (middle) and after (right) test stimulation (1 min, 6 Hz, 2 mA), administered following a 30 min stimulation (6 Hz, 2 mA). b. The somatosensory evoked potential (SEP) of test stimulation at baseline (0 min, BL) and after 30 and 60 minutes (blue, red, black, respectively, each averaged over 360 stimuli). c. SEP in response to test stimulation at baseline (blue) and following a 30 min stimulation (red). d. SEP in response to test stimulation at baseline (blue) and three time points following a 30 min stimulation (60 min, red; 90 min, green; 120 min, purple). e. Maximum amplitude of the SEP (absolute values) following a 30 min stimulation compared to baseline (n=6 rats, mean ± SEM, Wilcoxon, p=0.0312). f. Area under curve (AUC) of the SEP following a 30 min stimulation compared to baseline. (n=6 rats, mean ± SEM, Wilcoxon, p=0.0312). g. Top: 1-hour LFP trace from a representative rat. Bottom: 5 s magnifications of the above trace at selected time points (noted by asterisks). Left to right: baseline activity; during test stimulation; following cortical application of 0.1 mM albumin (Alb); during test stimulation post-Alb. h. SEP amplitude during test stimulation at baseline (normal aCSF, blue) and following 0.1 mM Alb (red). i. Maximum amplitude of the SEP during test stimulation post-Alb compared to baseline (n=5 rats, mean ± SEM, Wilcoxon, p=0.0312). j. AUC of the SEP post-Alb compared to baseline. (n=5 rats, mean ± SEM, Wilcoxon, p=0.0312). k. Power spectrum density of 10 min spontaneous activity before (blue) and post-Alb (red).

Revised version: a. Top: LFP trace from the rat L2/3 sensorimotor cortex before (left) during (middle) and after (right) test stimulation (1 min, 6 Hz, 2 mA, an excerpt of the stimulation trace is shown above the middle LFP trace). Bottom: LFP trace from the same rat before (left) during (middle) and after (right) test stimulation (1 min, 6 Hz, 2 mA), administered following a 30 min stimulation (6 Hz, 2 mA). b. The somatosensory evoked potential (SEP) of test stimulation at baseline (0 min, BL) and after 30 and 60 minutes (blue, red, black, respectively, each averaged over 360 stimuli). c. SEP in response to test stimulation at baseline (blue) and following a 30 min stimulation (red). d. SEP in response to test stimulation at baseline (blue) and three time points following a 30 min stimulation (60 min, red; 90 min, green; 120 min, purple). e. Maximum amplitude of the SEP (absolute values) following a 30 min stimulation compared to baseline (n=6 rats, mean ± SEM, Wilcoxon, p=0.0312). f. Area under curve (AUC) of the SEP following a 30 min stimulation compared to baseline. (n=6 rats, mean ± SEM, Wilcoxon, p=0.0312). g. Top: 1-hour LFP trace from a representative rat. Bottom: 5 s magnifications of the above trace at selected time points (noted by asterisks). Left to right: baseline activity; during test stimulation; following cortical application of 0.1 mM albumin (Alb); during test stimulation post-Alb. h. SEP amplitude during test stimulation at baseline (normal aCSF, blue) and following 0.1 mM Alb (red). i. Maximum amplitude of the SEP during test stimulation post-Alb compared to baseline (n=5 rats, mean ± SEM, Wilcoxon, p=0.0312). j. AUC of the SEP post-Alb compared to baseline. (n=5 rats, mean ± SEM, Wilcoxon, p=0.0312). k. Power spectrum density of 10 min spontaneous activity before (blue) and post-Alb (red) (p=0.0035, paired t-test).

Earlier version: To test whether stimulation-induced albumin extravasation has a role in LTP, we recorded local field potentials (LFP) before and after exposure of the cortex to 0.1 mM albumin (corresponding to 25% of serum concentration) (Ivens et al., 2007). Albumin application increased the amplitude and frequency of spontaneous activity (Figure 2g,k). Neuronal response to the stimulation showed synaptic potentiation in the presence of albumin compared to that recorded under aCSF (Figure 2h-j), indicating that cortical exposure to serum albumin results in synaptic potentiation.

Revised version: To test whether stimulation-induced albumin extravasation has a role in LTP, we recorded local field potentials (LFP) before and after exposure of the cortex to 0.1 mM albumin (corresponding to 25% of serum concentration) (Ivens et al., 2007). The relatively high concentration of albumin was chosen to account for factors that lower its effective tissue concentration, such as its low diffusion rate and its likelihood to encounter a degradation site or a cross-BBB efflux transporter (Tao and Nicholson, 1996; Zhang and Pardridge, 2001).

Added in the revised version: Albumin application increased the amplitude and frequency of spontaneous activity (Figure 2g,k). Neuronal response to the stimulation showed synaptic potentiation in the presence of albumin compared to that recorded under aCSF (Figure 2h-j), indicating that cortical exposure to serum albumin results in synaptic potentiation.

BBB modulation is activity-dependent

To test the neurotransmission pathways involved in regulation of BBB permeability during stimulation we exposed the cortical window to selective antagonists of ionotropic glutamate receptors (Figure 3a). The AMPA/kainate receptor blocker CNQX prevented the stimulation-induced BBB modulation, whereas the NMDAR blocker AP5 did not (Figure 3b-c). CNQX also reversibly blocked synaptic activity and prevented the vascular response to neuronal activation, indicated by the absence of vasodilation (AKA “neurovascular coupling”; Figure 3d-e, Figure S2a, S3a,c). These results indicate that stimulation-induced BBB modulation was mediated by neuronal activation. In contrast, AP5 did not affect baseline SEP (Figure 3d-e), nor the vascular response (Figure S2a, S3b). Synaptic potentiation was blocked under both CNQX or AP5 (Figure 3d-e, g-h, Figure S2a), consistent with the crucial role of AMPA and NMDA receptors in synaptic plasticity. Next, we tested for molecular correlates of synaptic strength by comparing postsynaptic density 95 (PSD-95) expression levels in the cortex at different time points. PSD-95 is known to increase following LTP and AMPAR trafficking (El-Husseini et al., 2000; Holtmaat and Svoboda, 2009). Western blot analysis showed significantly higher levels of PSD-95 expression in stimulated rats as compared to sham, at both 30 min and 24 h after stimulation (Figure S2c-d).

Stimulation-induced plasticity is associated with BBB modulation.

Earlier version: a. Timeline of the experimental protocol for stimulations and imaging with blockers application (CNQX/AP5 (50 μM)/mβCD (10 μM)/SJN (0.3 mM). b. NaFlu permeability maps of the cortical window before (control, left), and after CNQX/AP5/SJN + 30 min stimulation (right). c. PI following blockers application before and after stimulation compared to a 30 min stimulation. (% change from baseline, 30 min stim n=6, AP5 n=4, CNQX n=5, mβCD n=6, SJN n=5; mean ± SEM, Kruskal-Wallis with FDR correction, *p<0.05). d. SEP amplitude in response to test stimulation. baseline (blue); following application of CNQX (left, red) or AP5 (right, red); following CNQX + 30 min stimulation (left, black); and following AP5 + 30 min stimulation (right, black); e. Left: Max SEP amplitude to test stimulation following CNQX compared to baseline. (mean ± SEM, n=5, paired t-test, p=0.0176). Right: Max SEP amplitude following AP5 compared to baseline. f. SEP amplitude in response to test stimulation. baseline (blue); following SJN (black); following SJN + 30 min stimulation (red); and following mβCD (right). g. Max SEP amplitude following 30 min stimulation or albumin application compared to blockers, and blockers + 30 min stimulation. h. AUC of the SEP for test stimulation following 30 min stimulation or albumin application compared to blockers and blockers + 30 min stimulation. (g-h: % change from baseline, mean ± SEM, 30 min stim n=5, Alb n=4, AP5 n=5, CNQX n=5, SJN n=4, mβCD n=5; Kruskal-Wallis with FDR correction, *p<0.05).

Revised version: a. Timeline of the experimental protocol for stimulations and imaging with blockers application (CNQX/AP5 (50 μM)/mβCD (10 μM)/SJN (0.3 mM). b. NaFlu permeability maps of the cortical window before (control, left), and after CNQX/AP5/SJN + 30 min stimulation (right). Tracer accumulation area is marked with a white circle. c. PI following blockers application before and after stimulation compared to a 30 min stimulation. (% change from baseline, 30 min stim n=6, AP5 n=4, CNQX n=5, mβCD n=6, SJN n=5; mean ± SEM, Kruskal-Wallis with FDR correction, *p<0.05). d. SEP amplitude in response to test stimulation. baseline (blue); following application of CNQX (left, red) or AP5 (right, red); following CNQX + 30 min stimulation (left, black); and following AP5 + 30 min stimulation (right, black); e. Left: Max SEP amplitude to test stimulation following CNQX compared to baseline. (mean ± SEM, n=5, paired t-test, p=0.0176). Right: Max SEP amplitude following AP5 compared to baseline. f. SEP amplitude in response to test stimulation. baseline (blue); following SJN (black); following SJN + 30 min stimulation (red); and following mβCD (right). g. Max SEP amplitude following 30 min stimulation or albumin application compared to blockers, and blockers + 30 min stimulation. h. AUC of the SEP for test stimulation following 30 min stimulation or albumin application compared to blockers and blockers + 30 min stimulation. (g-h: % change from baseline, mean ± SEM, 30 min stim n=5, Alb n=4, AP5 n=5, CNQX n=5, SJN n=4, mβCD n=5; Kruskal-Wallis with FDR correction, *p<0.05).

Albumin extravasation and LTP require caveolae-mediated transcytosis

In endothelial cells, albumin is transported mainly via caveolae-mediated transcytosis (CMT) through binding to the gp60 receptor (Tiruppathi et al., 1997). To test whether CMT mediates stimulation-induced transcytosis of albumin (John et al., 2003) we perfused the exposed cortex with methyl-β-cyclodextrin (mβCD), which inhibits albumin uptake by disrupting caveolae (Schnitzer and Oh, 1994; Skotland et al., 2020). Cortical application of mβCD prevented stimulus-induced BBB opening (Figure 3c), without affecting vascular response to stimulation (Figure S3d). mβCD also prevented stimulus-induced LTP (Figure 3f-h), further supporting the role of CMT-mediated albumin extravasation in activity-dependent plasticity in vivo.

BBB modulation and LTP involve TGF-β signaling

Under pathological conditions, synaptogenesis and pathological plasticity that follow BBB dysfunction are mediated by the activation of TGF-βR1 signaling in astrocytes (Weissberg et al., 2015). TGF-β1 was also shown to directly induce BBB opening in endothelial cell cultures (McMillin et al., 2015), and to regulate hippocampal synaptic plasticity (Dahlmanns et al., 2022; Gradari et al., 2021). We therefore tested the effect of stimulation in the presence of the specific TGF-βR1 blocker SJN2511 (SJN, 0.3 mM (Gellibert et al., 2004; Weissberg et al., 2015)). Cortical exposure to SJN prevented stimulation-induced BBB opening (Figure 3b-c) with no effect on neurovascular coupling (Figure S3e). SJN also prevented stimulation-induced LTP (Figure 3f-h). These results demonstrate that TGF-β signaling is required for activity-dependent BBB modulation, and the consequent activity-dependent plasticity, but not for neurovascular coupling.

Regulation of BBB transport and plasticity genes

To explore gene-expression changes related to our stimulation protocol, we conducted bulk RNA-sequencing transcriptome analysis of tissue dissected from the sensorimotor area of the cortex contralateral and ipsilateral to stimulated limb. Tissue harvested 1 and 24 h post-stimulation (30 min) showed differentially expressed genes (DEGs) in the contralateral (“stimulated”) compared to the ipsilateral (“non-stimulated”) cortex (13.2% and 7.3% 24 and 1 h, respectively, Figure 4a, S4a). To assess the variability between paired samples of stimulated and non-stimulated cortex of each animal at two time points following stimulation (24 h vs 1 h) we used the Jensen–Shannon divergence metric (JSD) (Sudmant et al., 2015). JSD calculations using normalized counts indicated that the RNA expression 24 h after stimulation was significantly different from that of rats 1 h after stimulation (Figure S4b). Gene Ontology (GO) pathway enrichment analysis of DEGs from stimulated rats indicated DEGs were primarily involved in synaptic plasticity processes (Figure 4b). A significantly higher number of DEGs related to synaptic plasticity was found in the contralateral cortex of stimulated rats compared to ipsilateral (p<0.001, Figure 4d). No significant differences were found in BBB or inflammation related genes between the hemispheres, suggesting that the stimulation protocol is not associated with the transcriptional change observed in BBB dysfunction under pathological conditions (Cacheaux et al., 2009; Kim et al., 2017) (Figure 4d). In contrast, a larger amount of transporter genes such as solute carrier transporters (Slc) and ATP-binding cassette (ABCs) families were differentially expressed in the stimulated compared to the non-stimulated cortex (Figure S4c-d). Both DEG and enrichment analysis were consistent with activation of TGF-β signaling pathway (Figure 4d). Analysis of vascular cell-specific DEGs (Vanlandewijck et al., 2018) showed significantly more arterial smooth muscle cells specific genes in the stimulated cortex (Figure 4c, Figure S4e). These results support the roles of stimulation-induced modulation of BBB transcellular transport and TGF-β signaling in activity-dependent plasticity (Figure S5).

Neuronal activity regulates BBB transport and synaptic plasticity genes.

Earlier version: a. Scatter plot of gene expression from RNA-seq in the contralateral cortex 24 vs. 1 h after 30 min stimulation. The x axis represents the log fold change, and the y axis represents the mean expression levels. Blue dots indicate statistically significant differentially expressed genes (DEGs) by Wald Test (n=8 rats per group). b. Top Gene Ontologies (GO) enriched terms in the contralateral cortices of rats 24 vs. 1 h after stimulation. c. Vascular cell specific DEGs to pericytes (PC), arterial smooth muscle cells (aSMC), and venous smooth muscle cells (vSMC). (PC n=42, aSMC n=70, p=0.0017, Chi-square). d. Scatter plot of DEGs divided by groups of interest: BBB properties, NVU properties, Synaptic plasticity, and inflammatory related genes in the contralateral (red) vs. ipsilateral (black) cortices of stimulated rats. Circles represent genes with no significant differences between 1 and 24 h post-stimulation. Upward and downward triangles indicate significantly up- and down-regulated genes, respectively (synaptic plasticity n=53, ***p<0.001, Chi-square).

Revised version: a. Scatter plot of gene expression from RNA-seq in the contralateral somatosensory cortex 24 vs. 1 h after 30 min stimulation. The y axis represents the log fold change, and the x axis represents the mean expression levels (see methods, RNA Sequencing & Bioinformatics). Blue dots indicate statistically significant differentially expressed genes (DEGs) by Wald Test (n=8 rats per group). b. Top Gene Ontologies (GO) enriched terms in the contralateral cortices of rats 24 vs. 1 h after stimulation. c. Vascular cell specific DEGs to pericytes (PC), arterial smooth muscle cells (aSMC), and venous smooth muscle cells (vSMC). (PC n=42, aSMC n=70, p=0.0017, Chi-square). d. Scatter plot of DEGs divided by groups of interest: BBB properties, NVU properties, Synaptic plasticity, and inflammatory related genes in the contralateral (red) vs. ipsilateral (black) cortices of stimulated rats. Circles represent genes with no significant differences between 1 and 24 h post-stimulation. Upward and downward triangles indicate significantly up- and down- regulated genes, respectively. The contralateral somatosensory cortex was found to have a significantly higher number of DEGs related to synaptic plasticity, than the ipsilateral side (***p<0.001, Chi-square).

Physiological BBB modulation in the human brain

Finally, we tested for neuronal activity-induced BBB modulation in healthy human volunteers by measuring the blood-oxygenation-level dependent (BOLD) response to a 30 min stress-ball squeeze task and blood-to-brain transport using dynamic contrast-enhanced (DCE) MRI (Figure 5a-b). As expected, BOLD response was found in the pre- and post- central gyrus, corresponding with the primary motor and sensory cortices, respectively, of the hemisphere contralateral to the squeezing hand (Figure 5c-d). DCE-MRI showed higher BBB permeability to the gadolinium-based contrast agent gadoterate meglumine (Dotarem) in sensory/motor-related areas of subjects performing the task, as compared to controls (Figure 5e-g). Non activated cortical regions such as the medial frontal cortex and middle occipital gyrus did not differ in BBB permeability between task and controls (Figure 5g). Comparison of the spatial distribution of functional and permeability changes revealed a higher percentage of activated regions with modulated BBB in task performers than in control subjects (Figure 5e-g).

Cortical activation in fMRI co-localizes with BBB modulation.

Earlier version: a. Subjects were given an elastic stress-ball to squeeze continuously for the length of the session (30 min). b. Timeline of the experimental protocol for task performance and MRI sequences. c. Activation map for the localizer task displayed over the inflated brain of an exemplary subject. (p<0.05, Family wise error (FWE) corrected, neurologic convention). d. Activation maps for the localizer task displayed over anatomical axial slices of an exemplary subject. (p<0.05, FWE corrected, radiologic convention). e. Left: Superimposed masks of BBB modulated voxels (red) and fMRI activation (yellow). Right: Co-localized voxels map on the same slice (cyan). f. BBB modulation in the primary motor cortex (M1, precentral gyrus, PrG) and primary somatosensory cortex (S1, postcentral gyrus, PoG) for subjects performing the task compared to controls (% of voxels in each region with BBB leakage). (Task n=6, controls n=10, i–ipsi, c-contra, mean ± SEM, *p<0.05, two-way ANOVA with FDR correction). g. Heatmap of BBB modulated voxels percentage in motor/sensory related areas and some non-activated cortical regions of task vs. controls (+Task n=6, -Task n=10, i–ipsi, c-contra, mean ± SEM, *p<0.05, **p<0.001, ***p<0.0001, ns – non significant, two-way ANOVA with FDR correction).

Revised version: a. Subjects were given an elastic stress-ball to squeeze continuously for the length of the session (30 min). b. Timeline of the experimental protocol for task performance (the hand holding a ball represents the duration of the ball-squeezing task), and MRI sequences (T1-weighted, functional (fMRI), and dynamic contrast-enhanced (DCE) MRI). c. Activation map (voxels with statistically significant activation (t-statistic)) for the localizer task displayed over the inflated brain of an exemplary subject. (p<0.05, Family wise error (FWE) corrected, neurologic convention). d. Activation maps for the localizer task displayed over anatomical axial slices of an exemplary subject (p<0.05, FWE corrected, radiologic convention). White lines point to voxels of highest activation (t-statistic). e. Left: Superimposed masks of BBB modulated voxels (red) and fMRI activation (yellow). Right: Co-localized voxels map on the same slice (cyan). f. The percent of voxels with BBB leakage in the primary motor cortex (M1, precentral gyrus, PrG) and primary somatosensory cortex (S1, postcentral gyrus, PoG) for subjects performing the task compared to controls. (Task n=6, controls n=10, i–ipsi, c-contra, mean ± SEM, *p<0.05, two-way ANOVA with FDR correction). g. Heatmap of BBB modulated voxels percentage (as in panel f) in motor/sensory related areas and some non-activated cortical regions of task vs. controls (+Task n=6, -Task n=10, i–ipsi, c-contra, mean ± SEM, *p<0.05, **p<0.001, ***p<0.0001, ns – nonsignificant, two-way ANOVA with FDR correction).

Discussion

This study demonstrates that prolonged physiological stimulation modulates BBB permeability in the rat and human cortex. In rats, we show that prolonged paw stimulation results in increased cross-BBB extravasation of both low- and high- molecular weight tracers. We also show that this extravasation is: (1) focal and localized to the cortical area corresponding to the stimulation; (2) mediated by AMPA neurotransmission, TGF-β signaling, and caveolae transcytosis; and (3) associated with synaptic potentiation and increased postsynaptic density. In humans, we present the first evidence for increased BBB permeability in response to prolonged limb activity. Together, our results suggest that stimulation-induced BBB modulation may play a role in synaptic potentiation.

Physiological changes in cross-BBB transport have been previously reported in response to whisker stimulation, circadian rhythms, and age (Nishijima et al., 2010; Pulido et al., 2020; Yang et al., 2020). Here, we examined whether the mechanisms involved in this process may include those identified in our studies of pathological BBB modulation. Previous studies from our group showed that brain injury is associated with albumin transcytosis, albumin binding to TGF-βR in astrocytes, and pathological TGF-β-mediated plasticity (Salar et al., 2016). Here, we show that prolonged physiological sensory-motor activation is also associated with albumin extravasation (although at a significantly lower concentration than under pathological conditions, Figure 1l and Figure S1d) and TGF-βR signaling. We further show that blocking caveolae-mediated albumin transcytosis or TGF-βR signaling inhibits stimulation-evoked increase in BBB permeability and synaptic potentiation (without affecting the physiological hemodynamic response). These results suggest that caveolae albumin transcytosis and TGF-β signaling are key mediators of activity-dependent BBB modulation and subsequent synaptic plasticity. These findings are supported by reports linking TGF-β signaling to: synaptic plasticity induced by environmental enrichment (Dahlmanns et al., 2022); enhanced spine and synapse formation (Patel and Weaver, 2021); and hippocampal neuroplasticity (Gradari et al., 2021). Moreover, our finding of caveolae-mediated albumin transcytosis in pial arterioles is in line with previous studies showing: caveolae-dependent transcytosis of albumin (Tecedor et al., 2013); caveolae abundance in arteriolar ECs and not capillary ECs (Chow et al., 2020); and the contribution of caveolae transcytosis to neurovascular coupling (Andreone et al., 2017; Chow et al., 2020). Notably, in our experiments, caveolae inhibition prevented BBB modulation yet did not significantly affect neurovascular coupling, likely due to the contribution of other vasodilatory pathways such as endothelial nitric oxide synthase.

As mentioned above, albumin is a known activator of TGF-β signaling, and TGF-β has a well-established role in neuroplasticity. Interestingly, emerging evidence suggests that TGF-β also increases cross-BBB transcytosis (Betterton et al., 2022; Kaplan et al., 2020; McMillin et al., 2015; Schumacher et al., 2023). Hence, we propose the following two-part hypothesis for the TGF-β/BBB-mediated synaptic potentiation observed in our experiments: (1) prolonged stimulation triggers TGF-β signaling and increased caveolae-mediated transcytosis of albumin; and (2) extravasated albumin induces further TGF-β signaling, leading to synaptogenesis and additional cross-BBB transport – in a self-reinforcing positive feedback loop. Future research is needed to examine the validity of this hypothesis.

Our study also shows that physiological BBB modulation requires AMPAR glutamate signaling and not NMDA. This finding is consistent with studies showing the dominant role of AMPAR in electrophysiological and hemodynamic responses to somatosensory stimulation (Graves et al., 2021; Gsell et al., 2006; Zhang et al., 2015). Moreover, the well-established role of AMPAR in neuroplasticity, long-term potentiation (LTP), and experience-dependent sensory cortical plasticity (Daw et al., 1993; Feldman, 2009; Watt et al., 2004) further supports the plasticity-related role of activity-dependent BBB modulation.

Added in the revised version: Our gene analysis was designed to complement our in vivo and histological findings, by assessing the magnitude of change in DEGs. This analysis showed that: (1) the hemisphere contralateral to the stimulus has significantly more DEGs than the ipsilateral hemisphere; and (2) the DEGs were related to synaptic plasticity and TGF-β signaling. These findings strengthen the hypothesis raised by our experiments.

Earlier version: Our animal experiments show that a 30 min limb stimulation (at 6Hz and 2mA) increases cross-BBB influx, while a 1 min stimulation (of similar frequency and magnitude) does not. We believe that both types of stimulations fall within the physiological range because our continuous electrophysiological recordings showed no signs of epileptiform or otherwise pathological activity. Moreover, the recorded SEP levels were similar to those reported in previous physiological LTP studies in rats (Eckert and Abraham, 2010; Han et al., 2015; Mégevand et al., 2009) and humans (McGregor et al., 2016). Future studies are needed to explore the BBB modulating effects of additional stimulation protocols – with varying durations, frequencies, and magnitudes. Such studies may also elucidate the temporal and ultrastructural characteristics that differentiate between physiological and pathological BBB modulation.

Revised version: Our animal experiments show that a 30 min limb stimulation (at 6Hz and 2mA) increases cross-BBB influx, while a 1 min stimulation (of similar frequency and magnitude) does not. We argue that both types of stimulations fall within the physiological range because in rats, activity between 5-15 Hz represents physiological rhythmic whisker movement during environment exploration (Mégevand et al., 2009). Moreover, our continuous electrophysiological recordings showed no signs of epileptiform or otherwise pathological activity and the recorded SEP levels were similar to those reported in previous LTP studies in rats (Eckert and Abraham, 2010; Han et al., 2015; Mégevand et al., 2009) and humans (McGregor et al., 2016). In humans, skill acquisition often involves motor training sessions that last ≥30 minutes (Bengtsson et al., 2005; Classen et al., 1998) and result in physiological plasticity of sensory and motor systems (Classen et al., 1998; Draganski et al., 2004; Sagi et al., 2012). Hence, the experimental task in our human study (30 minutes of repetitive squeezing of an elastic stress-ball) is likely to represent physiological learning task involving neuronal activation in primarily motor and sensory areas (Halder et al., 2005). Future human and animal studies are needed to explore the BBB modulating effects of additional stimulation protocols – with varying durations, frequencies, and magnitudes. Such studies may also elucidate the temporal and ultrastructural characteristics that differentiate between physiological and pathological BBB modulation.

Removed in the revised version: In humans, skill acquisition often involves motor training sessions that last ≥30 minutes (Bengtsson et al., 2005; Classen et al., 1998) and result in physiological plasticity of sensory and motor systems (Classen et al., 1998; Draganski et al., 2004; Sagi et al., 2012). We, therefore, believe that the experimental task of our human study (30 minutes of repetitive squeezing of an elastic stress-ball) represents physiological activity, with neuronal activation in primarily motor and sensory areas (Halder et al., 2005). We further propose that the activity-related increase in BBB permeability we observed may be associated with learning.

Earlier version: A key limitation of our animal experiments is the fact they were performed under anesthesia, due to the complex nature of the experimental setup (i.e., simultaneous cortical imaging and electrophysiological recordings). Anesthetic agents can potentially alter neuronal activity, SEPs, CBF, and vascular responses (Aksenov et al., 2015; Lindauer et al., 1993; Masamoto and Kanno, 2012). To minimize these effects, we used ketamine-xylazine anesthesia, which unlike other anesthetics, was shown to maintain robust BOLD and SEP responses to neuronal activation (Franceschini et al., 2010; Shim et al., 2018). Another limitation of our animal study is the potentially non-specific effect of mβCD – an agent that disrupts caveola transport but may also lead to cholesterol depletion (Keller and Simons, 1998). To mitigate this issue, we used a very low mβCD concentration (10μM), orders of magnitude below the concentration reported to deplete cholesterol (Koudinov and Koudinova, 2001). Lastly, our animal study is limited by the inclusion of solely male rats. While our findings in humans did not point to sex-related differences in stimulation-evoked BBB modulation, larger animals and human studies are needed to examine this question.

Revised version: A key limitation of our animal experiments is the fact they were performed under anesthesia, due to the complex nature of the experimental setup (i.e., simultaneous cortical imaging and electrophysiological recordings). Anesthetic agents can potentially alter neuronal activity, SEPs, CBF, and vascular responses (Aksenov et al., 2015; Lindauer et al., 1993; Masamoto and Kanno, 2012). To minimize these effects, we used ketamine-xylazine anesthesia, which unlike other anesthetics, was shown to maintain robust BOLD and SEP responses to neuronal activation (Franceschini et al., 2010; Shim et al., 2018). Notably, the antagonistic effect of ketamine on NMDA receptors might attenuate the magnitude of SEP potentiation recorded in our experiments (Anis et al., 1983; Salt et al., 1988).

Added in the revised version: Another limitation of our animal study is the potentially non-specific effect of mβCD – an agent that disrupts caveola transport but may also lead to cholesterol depletion (Keller and Simons, 1998). To mitigate this issue, we used a very low mβCD concentration (10μM), orders of magnitude below the concentration reported to deplete cholesterol (Koudinov and Koudinova, 2001). Lastly, our animal study is limited by the inclusion of solely male rats. While our findings in humans did not point to sex-related differences in stimulation-evoked BBB modulation, larger animals and human studies are needed to examine this question.

To conclude, our study suggests that BBB modulation in response to physiological neuronal activity involves mechanisms previously identified in pathology-induced BBB changes (despite the significantly smaller increase in cross-BBB influx under physiological conditions). Our findings further suggest that physiological modulation of cross-BBB influx may play a role in synaptic plasticity, calling for future research into the dynamics of BBB function in health and disease.

Methods

Added in the revised version: Study design

Added in the revised version: This study included rat and human experiments. The rat experiments were designed to explore the mechanisms involved in altered BBB permeability following physiological stimulation and the potential link to synaptic plasticity. The experimental paradigm involved a 30-min rat limb stimulation at a frequency comparable to rhythmic whisker activity during environment exploration (Mégevand et al., 2009). The effects of stimulation on BBB permeability were studied using in vivo and histopathological imaging. The effects of stimulation on synaptic plasticity were evaluated by examining the somatosensory evoked potentials in response to a 1-min test stimulation before and after the 30-min stimulation. A sub-set of experiments involved the application of blockers to test their effect on BBB permeability and synaptic plasticity. Another sub-set of experiments was designed to test the effects of stimulation on gene expression.

Added in the revised version: In humans, the experimental paradigm involved a 30-min session of squeezing a stress ball. fMRI was used to localize the brain area activated during this motor task and dynamic contrast-enhanced (DCE) MRI was used to evaluate BBB permeability post stimulation.

Added in the revised version: All experiments and data analysis were performed blindly to the treatment group.

Animal preparation and surgical procedures

All experimental procedures were approved by the Animal Care and Use Ethical Committee at the Ben-Gurion University of the Negev, Beer-Sheva, Israel.

Adult Sprague Dawley male rats (300-350 g; Harlan Laboratories) were kept under a 12:12 h light and dark regimen and supplied with drinking water and food ad libitum. Surgical procedures were performed as reported previously (Prager et al., 2010). Rats were deeply anesthetized by intraperitoneal administration of ketamine (100 mg/ml, 0.08 ml/100 g) and xylazine (20 mg/ml, 0.06 ml/100 g). The tail vein was cannulated, and animals were placed in a stereotactic frame under a SteREO Lumar V12 fluorescence microscope (Zeiss). Body temperature was continuously monitored and kept stable at 37 ± 0.5°C using a feedback-controlled heating pad (Physitemp). Heart rate, breath rate, and oxygen saturation levels were continuously monitored using MouseOx (STARR Life Sciences). A cranial section (4 mm caudal, 2 mm frontal, 5 mm lateral to bregma) was removed over the right sensorimotor cortex. The dura and arachnoid layers were removed, and the exposed cortex was continuously perfused with artificial CSF (aCSF (Prager et al., 2010), containing (in mM): 124 NaCl, 26 NaHCO3, 1.25 NaH2PO4, 2 MgSO4, 2 CaCl2, 3 KCl, and 10 glucose, pH 7.4).

Stimulation protocol

The left forelimb or hind limb of the rat was stimulated using Isolated Stimulator device (AD Instruments) attached with two subdermal needle electrodes (0.1 ms square pulses, 2-3 mA) at 6 Hz frequency. Test stimulation consisted of 360 pulses (60 s) and delivered before (as baseline) and after long-duration stimulation (30 min, referred throughout the text as ‘stimulation’). In control and albumin rats, only short-duration stimulations were performed. Under sham stimulation, electrodes were placed without delivering current.

Electrophysiology

In-vivo recordings

Earlier version: Stimulus-evoked potentials were recorded in the somatosensory cortex. Local field- and somatosensory evoked potentials (LFP, SEP) were recorded using a glass microelectrode (1.5 × 1.1 mm, 1-2 µm tip, 2-10 MΩ) filled with aCSF, which was inserted with an isolated, chlorinated silver wire connected to a headstage and amplifier (EXT-02B, npi electronic, Germany). The electrode was attached to a digital micromanipulator (MP-225, Sutter Instruments, CA, USA) and was carefully placed over the identified area of the sensorimotor response to the stimulation.

Revised version: Stimulus-evoked potentials were recorded in the somatosensory cortex. Local field- and somatosensory evoked potentials (LFP, SEP) were recorded using a glass microelectrode (1.5 x 1.1 mm, 1-2 µm tip, 2-10 MΩ) filled with aCSF, which was inserted with an isolated, chlorinated silver wire connected to a headstage and amplifier (EXT-02B, npi electronic, Germany). The electrode was attached to a digital micromanipulator (MP-225, Sutter Instruments, CA, USA) and was carefully placed over the identified area of the sensorimotor response to the stimulation.

Analysis of electrophysiological recordings

Earlier version: Data acquisition was performed using PowerLab and LabChart (AD Instruments), for stimulation, and extracellular recordings. Signals were digitized (1 kHz) and filtered (high-pass 1 Hz, low pass 45 or 300 Hz). Analyses were performed using in house MATLAB scripts. Two measures were used to quantify synaptic plasticity: the maximal amplitude of the dominant negative peak and the absolute area under curve (AUC) of the SEP waveform (i.e., all positive and negative peaks between 10–160 ms post-stimulus period). For power analysis in albumin experiments, 10 min segments recorded before starting and after terminating albumin perfusion were used.

Revised version: Data acquisition was performed using PowerLab and LabChart (AD Instruments), for stimulation, and extracellular recordings. Signals were digitized (1 kHz) and filtered (high-pass 1 Hz, low pass 45 or 300 Hz). Analyses were performed using in house MATLAB scripts. Synaptic plasticity was then quantified using the following two measures: maximal amplitude of the dominant negative peak and the absolute area under curve (AUC) of the SEP waveform (i.e., all positive and negative peaks between 10–160 ms post-stimulus period). To overcome potential variability in SEP morphology between animals (Mégevand et al., 2009), each animal’s plasticity measures (max amplitude and AUC of post stimulation SEP) were compared to the same measures at baseline.

Added in the revised version: For power analysis in albumin experiments, 10 min segments recorded before starting and after terminating albumin perfusion were used.

Drug Application

Earlier version: In some experiments the following blockers were added to the aCSF and continuously perfused to the cortical window for 50 min: D-(-)-2-Amino-5-phosphonopentanoic acid (AP5, 50 μM), 6-cyano-7-nitroquinoxaline-2,3-dione (CNQX, 50 μM), methyl-β-cyclodextrin (MβCD, 10 μM), or SJN2511 (SJN, 0.3 mM (Weissberg et al., 2015)). To simulate BBB opening, bovine serum albumin (Alb, 0.1 mM (Seiffert et al., 2004)) was added to the aCSF and continuously perfused to the cortical window for 30 min.

Revised version: In some experiments the following blockers were added to the aCSF and continuously perfused to the cortical window for 50 min: D-(-)-2-Amino-5-phosphonopentanoic acid (AP5, 50 μM), 6-cyano-7-nitroquinoxaline-2,3-dione (CNQX, 50 μM), methyl-β-cyclodextrin (MβCD, 10 μM), or SJN2511 (SJN, 0.3 mM (Weissberg et al., 2015)). To simulate BBB opening, bovine serum albumin (Alb, 0.1 mM (Seiffert et al., 2004)) was added to the aCSF and continuously perfused to the cortical window for 30 min. Experimenter was blind to the selected compound.

In-vivo imaging

Earlier version: Dynamic imaging of regional cerebral blood flow, optical imaging of cortical oxygenation, neurovascular coupling (NVC), and vessel diameter measurements were carried out as previously described (Bouchard et al., 2009; Levi et al., 2012; Prager et al., 2010). Before, during, and after electrical limb stimulation, full-resolution (2560 × 2160 px) images were obtained for regional localization of the hemodynamic response, followed by acquisition of localized (512 × 512 px) images of cortical surface vessels (2 frames/s, EMCCD camera, Neo DC-152Q-COO-FI; Andor Technology) under 525/50 band pass filter (38 HE, Zeiss). Images were analyzed offline using in-house MATLAB (MathWorks) scripts.

Revised version: Dynamic imaging of regional cerebral blood flow, optical imaging of cortical oxygenation, neurovascular coupling (NVC), and vessel diameter measurements were carried out as previously described (Bouchard et al., 2009; Levi et al., 2012; Prager et al., 2010). Before, during, and after electrical limb stimulation, full-resolution (2560 x 2160 px) images were obtained for regional localization of the hemodynamic response, followed by acquisition of localized (512 × 512 px) images of cortical surface vessels (2 frames/s, EMCCD camera, Neo DC-152Q-COO-FI; Andor Technology) under 525/50 band pass filter (38 HE, Zeiss). Images were analyzed offline using in-house MATLAB (MathWorks) scripts.

Earlier version: Fluorescent angiography and BBB permeability measurements and analyses were performed as described previously (Prager et al., 2010; Vazana et al., 2016). The non-BBB permeable fluorescent dye sodium fluorescein (NaFlu) was injected intravenously (0.2 ml, 1 mg/ml). Images of cortical surface vessels were obtained before, during, and after tracer injection at 5 frames per second. Across all experiments, acquired images were the same size (512 × 512 pixel, ∼1×1 mm), centered above the responding arteriole. Images were analyzed offline using MATLAB as described (Vazana et al., 2016). Briefly, image registration and segmentation were performed to produce a binary image, separating blood vessels from extravascular regions. For each extravascular pixel, a time curve of signal intensity over time was constructed. To determine whether an extravascular pixel had tracer accumulation over time (due to BBB permeability), the pixel’s intensity curve was divided by that of the responding artery (i.e., the arterial input function, AIF, representing tracer input). This ratio was termed the BBB permeability index (PI), and extravascular pixels with PI > 1 were identified as pixels with tracer accumulation due to BBB permeability.

Revised version: Fluorescent angiography and BBB permeability measurements and analyses were performed as described previously (Prager et al., 2010; Vazana et al., 2016). The non-BBB permeable fluorescent dye sodium fluorescein (NaFlu) was injected intravenously (0.2 ml, 1 mg/ml). Images of cortical surface vessels were obtained before, during, and after tracer injection at 5 frames per second. Across all experiments, acquired images were the same size (512 × 512 pixel, ∼1x1 mm), centered above the responding arteriole. Images were analyzed offline using MATLAB as described (Vazana et al., 2016). Briefly, image registration and segmentation were performed to produce a binary image, separating blood vessels from extravascular regions. For each extravascular pixel, a time curve of signal intensity over time was constructed. To determine whether an extravascular pixel had tracer accumulation over time (due to BBB permeability), the pixel’s intensity curve was divided by that of the responding artery (i.e., the arterial input function, AIF, representing tracer input). This ratio was termed the BBB permeability index (PI), and extravascular pixels with PI > 1 were identified as pixels with tracer accumulation due to BBB permeability.

Histology and immunohistochemistry

Earlier version: Albumin extravasation was confirmed histologically in separate cohorts of rats that were anesthetized and stimulated without craniotomy surgery. Assessment of albumin extravasation was performed using a well-established approach that involves peripheral injection of either labeled-albumin (bovine serum albumin conjugated to Alexa Flour 488, Alexa488-Alb) or albumin-labeling dye (Evans blue, EB – a dye that binds to endogenous albumin and forms a fluorescent complex), followed by histological analysis of brain tissue (Ivens et al., 2007; Obermeier et al., 2013 ; Lapilover et al., 2012; Veksler et al., 2020, Ahishali et al., 2020). Since extravasated albumin is taken up by astrocytes (Ivens et al., 2007; Obermeier et al., 2013 and others), it can be visualized in the brain neuropil after brain removal and fixation (Ivens et al., 2007; Lapilover et al., 2012; Veksler et al., 2020, Ahishali et al., 2020). Five rats were injected with Alexa488-Alb (1.7 mg/ml) and five with EB (2%, 20 mg/ml, n=5). The injections were administered via the tail vein. Following injection, rats were transcardially perfused with cold phosphate-buffered saline (PBS), followed by PBS containing 4% paraformaldehyde (PFA). Brains were then removed, fixed overnight (4% PFA, 4°C), cryoprotected with sucrose gradient (10% followed by 20 and 30% sucrose in PBS) and frozen in optimal cutting temperature (OCT) compound. Coronal sections (40-μm thick) were obtained using a freezing microtome (Leica Biosystems) and imaged for dye extravasation using a fluorescence microscope (Axioskop 2; Zeiss) equipped with a CCD digital camera (AxioCam MRc 5; Zeiss).

Revised version: Albumin extravasation was confirmed histologically in separate cohorts of rats that were anesthetized and stimulated without craniotomy surgery. Assessment of albumin extravasation was performed using a well-established approach that involves peripheral injection of either labeled-albumin (bovine serum albumin conjugated to Alexa Flour 488, Alexa488-Alb) or albumin-labeling dye (Evans blue, EB – a dye that binds to endogenous albumin and forms a fluorescent complex), followed by histological analysis of brain tissue (Ahishali and Kaya, 2020; Ivens et al., 2007; Lapilover et al., 2012; Obermeier et al., 2013; Veksler et al., 2020). Since extravasated albumin is taken up by astrocytes (Ivens et al., 2007; Obermeier et al., 2013, and others), it can be visualized in the brain neuropil after brain removal and fixation (Ahishali and Kaya, 2020; Ivens et al., 2007; Lapilover et al., 2012; Veksler et al., 2020). Five rats were injected with Alexa488-Alb (1.7 mg/ml) and five with EB (2%, 20 mg/ml, n=5). The injections were administered via the tail vein. Following injection, rats were transcardially perfused with cold phosphate-buffered saline (PBS), followed by PBS containing 4% paraformaldehyde (PFA). Brains were then removed, fixed overnight (4% PFA, 4°C), cryoprotected with sucrose gradient (10% followed by 20 and 30% sucrose in PBS) and frozen in optimal cutting temperature (OCT) compound. Coronal sections (40-μm thick) were obtained using a freezing microtome (Leica Biosystems) and imaged for dye extravasation using a fluorescence microscope (Axioskop 2; Zeiss) equipped with a CCD digital camera (AxioCam MRc 5; Zeiss).

For immunostaining, brain slices were incubated in blocking solution (5% donkey serum in 0.1% Triton X-100/tris-buffered saline (TBS)), then incubated in primary antibody, followed by secondary antibody. Finally, slices were mounted and stained for DAPI (DAPI-Fluoromount-G, Invitrogen) to label nuclei. Staining was performed against glial fibrillary acidic protein (mouse anti GFAP−Cy3™, Sigma-Aldrich, C9205), Secondary antibody was anti-mouse Alexa Fluor 488 (Thermo-Fisher, A21206).

Immunoassays

Enzyme-linked immunosorbent assay (ELISA) for albumin extravasation

To assess the degree of albumin extravasation following stimulation, forty adult Sprague Dawley male rats (200-300g) were randomly divided into five groups. Rats were anesthetized (ketamine (75 mg/kg) and xylazine (5 mg/kg)) and underwent either a 30 min stimulation, sham stimulation (electrodes placed without current delivery), or photothrombosis stroke (PT) as previously described (Lippmann et al., 2017; Schoknecht et al., 2014). Briefly, Rose Bengal was administered intravenously (20 mg/kg) and a halogen light beam was directed for 15 min onto the intact exposed skull over the right somatosensory cortex. Then, rats were deeply anaesthetized and transcardially perfused with cold PBS, and brains were extracted at three time points following stimulation (immediately (real and sham), 4 and 24 h post stimulation) and 24 h post stroke. The ipsi- and contra lateral somatosensory cortices were dissected, weighed, and snap-frozen on dry ice. Fresh frozen samples were then homogenized in PBS + 1% Triton X-100 (v/v). Tissue homogenates were centrifuged at 13000 g for 10 min and supernatants were stored at −80°C until further analysis. Albumin levels were determined in diluted samples, using the Rat Albumin ELISA kit (Bethyl Laboratories, TX, USA) according to the manufacturer’s instructions.

Western Blotting

Earlier version: Brain tissues were collected from stimulated and control rats as described above. Protein lysates were separated using 10% sodium dodecyl sulfate polyacrylamide gel electrophoresis (SDS–PAGE, Bio-Rad) and then transferred to nitrocellulose membranes (Bio-Rad). Membranes were washed with TBS + 0.1% Tween20 (TBST, 3 × 5 min), then blocked with 5% BSA in TBS for 1 h. Membranes were then incubated overnight at 4°C with primary antibodies, followed by 2 h incubation with secondary antibodies. Protein bands were visualized with the enhanced chemiluminescent HRP substrate Crescendo Immobilon (Millipore). Analysis was performed in ImageJ, and band density levels were expressed as fold-changes to controls after normalization to β-actin and controls within the same blot. Primary antibodies used were mouse anti-PSD-95 (1:1500, Thermo Fisher Scientific, MA1-045), anti-β-actin (Loading control, 1:2000, Abcam, ab8226). Secondary antibodies were Goat anti-mouse (1:4000, Abcam, ab205719) and Goat anti-rabbit (1:4000, Abcam, ab97051).

Revised version: Brain tissues were collected from stimulated and control rats as described above. Protein lysates were separated using 10% sodium dodecyl sulfate polyacrylamide gel electrophoresis (SDS–PAGE, Bio-Rad) and then transferred to nitrocellulose membranes (Bio-Rad). Membranes were washed with TBS + 0.1% Tween20 (TBST, 3 x 5 min), then blocked with 5% BSA in TBS for 1 h. Membranes were then incubated overnight at 4°C with primary antibodies, followed by 2 h incubation with secondary antibodies. Protein bands were visualized with the enhanced chemiluminescent HRP substrate Crescendo Immobilon (Millipore). Analysis was performed in ImageJ, and band density levels were expressed as fold-changes to controls after normalization to β-actin and controls within the same blot. Primary antibodies used were mouse anti-PSD-95 (1:1500, Thermo Fisher Scientific, MA1-045), anti-β-actin (Loading control, 1:2000, Abcam, ab8226). Secondary antibodies were Goat anti-mouse (1:4000, Abcam, ab205719) and Goat anti-rabbit (1:4000, Abcam, ab97051).

RNA Extraction & Purification

Total RNA was extracted from 100-150 mg of flash-frozen rat somatosensory cortex tissue using TRIzol reagent (ThermoFisher Scientific #15596026) as previously described (Rio et al., 2010). Glycoblue (Invitrogen #AM9515) was added to assist with RNA pellet isolation. Pellets were washed with 75% ethanol and allowed to dry before being dissolved in 20 μL of RNase-DNase free water. 2 μL of 10x DNase I Buffer (New England Biolabs, #m0303s), 1 μL of rDNase I (ThermoFisher Scientific, AM2235), and DNase Inactivation Reagent (ThermoFisher Scientific, AM1907) were added to each RNA sample for purification.

RNA Sequencing & Bioinformatics

Earlier version: Library preparation was performed by the QB3-Berkeley Functional Genomics Laboratory at UC Berkeley. mRNA selection was done with the oligo-dT beads. Both total RNA and oligo dT bead selected mRNA quality was assessed on an Agilent 2100 Bioanalyzer. Libraries were prepared using the KAPA RNA Hyper Prep kit (Roche KK8540). Truncated universal stub adapters were ligated to cDNA fragments, which were then extended via PCR using unique dual indexing primers into full length Illumina adapters. Library quality was checked on an AATI (now Agilent) Fragment Analyzer. Libraries were then transferred to the QB3-Berkeley Vincent J. Coates Genomics Sequencing Laboratory, also at UC Berkeley. Library molarity was measured via quantitative PCR with the KAPA Library Quantification Kit (Roche KK4824) on a BioRad CFX Connect thermal cycler. Libraries were then pooled by molarity and sequenced on an Illumina NovaSeq 6000 S4 flowcell for 2 × 150 cycles, targeting at least 25M reads per sample. Fastq files were generated and demultiplexed using Illumina bcl2fastq2 v2.20. No other processing, trimming, or filtering was conducted. RNA-seq reads were aligned using STAR v2.9.1a against UCSC rat genome rn6. Downstream analysis was performed using FastQC v0.11.9, MultiQC v1.11, and featureCounts. R studio was used for data analysis, R v4.0.2. with the packages tidyverse v1.3.1, IHW, DESeq2. Gene ontology (GO) analysis of enriched genes in ipsi- and contra-lateral somatosensory cortices of stimulated rats was performed using Metascape (Zhou et al., 2019) (http://metascape.org/). Analysis of cell-specific and vascular zonation genes was performed as described (Vanlandewijck et al., 2018), using the database provided in (http://betsholtzlab.org/VascularSingleCells/database.html).

Revised version: Library preparation was performed by the QB3-Berkeley Functional Genomics Laboratory at UC Berkeley. mRNA selection was done with the oligo-dT beads. Both total RNA and oligo dT bead selected mRNA quality was assessed on an Agilent 2100 Bioanalyzer. Libraries were prepared using the KAPA RNA Hyper Prep kit (Roche KK8540). Truncated universal stub adapters were ligated to cDNA fragments, which were then extended via PCR using unique dual indexing primers into full length Illumina adapters. Library quality was checked on an AATI (now Agilent) Fragment Analyzer. Libraries were then transferred to the QB3-Berkeley Vincent J. Coates Genomics Sequencing Laboratory, also at UC Berkeley. Library molarity was measured via quantitative PCR with the KAPA Library Quantification Kit (Roche KK4824) on a BioRad CFX Connect thermal cycler. Libraries were then pooled by molarity and sequenced on an Illumina NovaSeq 6000 S4 flowcell for 2 x 150 cycles, targeting at least 25M reads per sample. Fastq files were generated and demultiplexed using Illumina bcl2fastq2 v2.20. No other processing, trimming, or filtering was conducted. RNA-seq reads were aligned using STAR v2.9.1a against UCSC rat genome rn6. Downstream analysis was performed using FastQC v0.11.9, MultiQC v1.11, and featureCounts. R studio was used for data analysis, R v4.0.2. with the packages tidyverse v1.3.1, IHW, DESeq2. Gene ontology (GO) analysis of enriched genes in ipsi- and contra-lateral somatosensory cortices of stimulated rats was performed using Metascape (Zhou et al., 2019) (http://metascape.org/). Analysis of cell-specific and vascular zonation genes was performed as described (Vanlandewijck et al., 2018), using the database provided in (http://betsholtzlab.org/VascularSingleCells/database.html).

Added in the revised version: Human study

Removed in the revised version: Magnetic Resonance Imaging

Participants

Male and female healthy individuals, aged 18-35, with no known neurological or psychiatric disorders were recruited to undergo MRI scanning while performing a motor task (n=6; 3 males and 3 females). MRI scans of 10 sex- and age- matched individuals (with no known neurological or psychiatric disorders) who did not perform the task were used as control data (n=10; 5 males and 5 females).

All procedures were approved by the Soroka University Medical Center Institutional Review Board. All participants gave their written informed consent before participation.

Added in the revised version: Magnetic Resonance Imaging

Removed in the revised version: Data acquisition

Earlier version: MRI scans were performed in the imaging center at the Soroka Medical Center using a 3T Philips Ingenia MRI scanner. A high-resolution T1-weighted anatomical scan (3D gradient echo, TE/TR 3.7/8.2 ms, voxel size of 1 mm3, 256 × 256 acquisition matrix), followed by a T2-weighted scan (TE/TR = 90/3000 ms, voxel size 0.45 × 0.45 × 4 mm). Functional data were collected by using a gradient echo EPI, with voxel size of 3 × 3 × 3 mm, TE/TR = 35/2000 ms, 83 × 80 acquisition matrix, flip angle of 90°, and FOV 109 mm. For calculation of pre-contrast longitudinal relaxation times, variable flip angle method was used (3D T1w-FFE, TE/TR = 2/10 ms, acquisition matrix: 256 × 256, voxel size: 0.89 × 0.89 × 6 mm, flip angles: 5, 15, 20 and 25°). Dynamic contrast-enhanced (DCE) sequence was then acquired (Axial, 3D T1w-FFE, TE/TR = 2/4 ms, acquisition matrix: 192 × 187 (reconstructed to 256 × 256), voxel size: 0.9 × 0.9 × 6 mm, flip angle: 20°, ýt = 10 s, temporal repetitions: 100, total scan length: 16.7 min). An intravenous bolus injection (0.1 mmol/kg, 0.5 M) of the gadolinium-based contrast agent gadoterate meglumine (Dotarem, Guerbet, France) was administered using an automatic injector after 5 dynamic scans at a rate of 1.5 ml/s.

Revised version: MRI scans were performed in the imaging center at the Soroka Medical Center using a 3T Philips Ingenia MRI scanner. A high-resolution T1-weighted anatomical scan (3D gradient echo, TE/TR 3.7/8.2 ms, voxel size of 1 mm3, 256 x 256 acquisition matrix), followed by a T2-weighted scan (TE/TR = 90/3000 ms, voxel size 0.45 x 0.45 x 4 mm). Functional data were collected by using a gradient echo EPI, with voxel size of 3 x 3 x 3 mm, TE/TR = 35/2000 ms, 83 x 80 acquisition matrix, flip angle of 90°, and FOV 109 mm. For calculation of pre-contrast longitudinal relaxation times, variable flip angle method was used (3D T1w-FFE, TE/TR = 2/10 ms, acquisition matrix: 256 x 256, voxel size: 0.89 x 0.89 x 6 mm, flip angles: 5, 15, 20 and 25°). Dynamic contrast-enhanced (DCE) sequence was then acquired (Axial, 3D T1w-FFE, TE/TR = 2/4 ms, acquisition matrix: 192 x 187 (reconstructed to 256 x 256), voxel size: 0.9 x 0.9 x 6 mm, flip angle: 20°, Δt = 10 s, temporal repetitions: 100, total scan length: 16.7 min). An intravenous bolus injection (0.1 mmol/kg, 0.5 M) of the gadolinium-based contrast agent gadoterate meglumine (Dotarem, Guerbet, France) was administered using an automatic injector after 5 dynamic scans at a rate of 1.5 ml/s.

Added in the revised version: fMRI Localizer Motor Task

Added in the revised version: fMRI was used to localize the brain area activated in response to the experimental motor task. During the fMRI scan, participants were asked to squeeze the stress-ball on cue (with green screen indicating Go and red screen indicating Stop) for a total of 5 experimental blocks. Task-related activation was evaluated for each cue, and voxels with statistically significant activation (t-statistic) were used to construct an ‘activation map’. The statistical threshold for clusters was set at p<0.05 and corrected for multiple comparisons using family-wise error (FWE) over the whole brain (uncorrected threshold p<0.001) to minimize detection of false positives (type I error).

Preprocessing of functional data

fMRI data preprocessing was performed using the Statistical Parametric Mapping package (SPM12; Wellcome Trust Centre for Neuroimaging, London, UK) for MATLAB. The first ten functional images of each run series were discarded to allow stabilization of the magnet. Functional images were realigned and co-registered to the T1-weighted structural image, followed by segmentation and normalization to the Montreal Neurological Institute (MNI) space. Spatial smoothing using Gaussian kernel (FWHM of 6 mm) was performed.

Removed in the revised version: fMRI Localizer Motor Task

Removed in the revised version: fMRI was used to localize the brain area activated in response to the experimental motor task. During the fMRI scan, participants were asked to squeeze the stress-ball on cue (with green screen indicating Go and red screen indicating Stop) for a total of 5 experimental blocks. Task-related activation was evaluated for each cue, and voxels with statistically significant activation (t-statistic) were used to construct an ‘activation map’. The statistical threshold for clusters was set at p<0.05 and corrected for multiple comparisons using family-wise error (FWE) over the whole brain (uncorrected threshold p<0.001) to minimize detection of false positives (type I error).

BBB permeability quantification

Analysis was performed as reported (Serlin et al., 2019; Veksler et al., 2020). Briefly, preprocessing included image registration and normalization to MNI space using SPM12. BBB permeability was calculated using in-house MATLAB script. A linear fit was applied to the later stage of the scan and the slope of contrast agent concentration changes over time was calculated for each voxel. Positive slopes reflected contrast agent accumulation due to BBB modulation. To compensate for physiological (e.g., heart rate, blood flow) and technical (e.g., injection rate) variability, slopes were normalized to the slope of each subject’s transverse sinus. Values were considered to reflect modulated BBB when exceeding the 95th percentile of the corresponding mean cumulative distribution function from healthy young controls. Voxels that reflect modulated BBB values were color-coded and superimposed on the anatomical scans to illustrate the kinetics, localization, and extent of BBB pathology.

Statistical analysis

Earlier version: All statistical tests were performed using Prism9 (GraphPad Software) and are indicated in the figure legends. Tests were two-tailed and corrected for multiple comparisons with two-stage linear step-up false discovery rate (FDR) procedure of Benjamini, Krieger and Yekutieli whenever applicable. Paired tests were used for inter-hemispheric comparisons. Non-parametric tests were used for data which was not normally distributed. Differences in fluorescence in Figure 1j-k were tested using a nested t-test. Differences in BBB permeability index, area under the curve and maximum amplitude of the somatosensory evoked potential before and after stimulation, were tested using the Wilcoxon matched-pairs signed rank test. Differences between groups in ELISA, area under the curve and maximum amplitude of the SEP were analyzed using Kruskal-Wallis test with FDR correction. Differences in the amount of DEGs in Figures 4c-d, S4c-e were analyzed using chi-square contingency test. Differences in BBB permeability across regions in the human brain were analyzed using two-way ANOVA with FDR correction. P≤ 0.05 was defined as the statistical significance level.

Revised version: All statistical tests were performed using Prism9 (GraphPad Software) and are indicated in the figure legends. Tests were two-tailed and corrected for multiple comparisons with two-stage linear step-up false discovery rate (FDR) procedure of Benjamini, Krieger and Yekutieli whenever applicable. Paired tests were used for inter-hemispheric comparisons. Non-parametric tests were used for data which was not normally distributed. Differences in fluorescence in Figure 1j-k were tested using a nested t-test. Differences in BBB permeability index, area under the curve and maximum amplitude of the somatosensory evoked potential before and after stimulation, were tested using the Wilcoxon matched-pairs signed rank test. Differences between groups in ELISA, area under the curve and maximum amplitude of the SEP were analyzed using Kruskal-Wallis test with FDR correction. Differences in the amount of DEGs in Figures 4c-d, S4c-e were analyzed using chi-square contingency test. Differences in BBB permeability across regions in the human brain were analyzed using two-way ANOVA with FDR correction. P≤ 0.05 was defined as the statistical significance level. Analyses were blind to experimental conditions.

Supporting information